---
evaluation_interval: 1m
rule_files:
  - alerts_rules.yml
  - mentions.yml
tests:
  - interval: 10s
    input_series:
      - series: ethereum_validators_monitoring_user_validators{nos_name="Operator 1", status="slashed"}
        values: 0+0x48 1+0x47 2+0x47 2+0x47 2+0x47 0+0x47 2+0x47 3+0x47
      - series: ethereum_validators_monitoring_user_validators{nos_name="Operator 2", status="slashed"}
        values: 0+0x48 1+0x47 2+0x47 3+0x47 3+0x47 4+0x47 0+0x47 0+0x47
      - series: ethereum_validators_monitoring_user_validators{nos_name="Operator 3", status="slashed"}
        values: 0+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 0+0x47 0+0x47
      - series: ethereum_validators_monitoring_user_validators{nos_name="Operator 4", status="slashed"}
        values: 0+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 0+0x47 0+0x47
      - series: ethereum_validators_monitoring_user_validators{nos_name="Operator 5", status="slashed"}
        values: 0+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 0+0x47 0+0x47
      # - series: ethereum_validators_monitoring_mentions{nos_global_index="1;0", mentions="op_0:ping_1,op_0:ping_2", alerts="UserSlashedValidators"}
      #   values: 1+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47
      # - series: ethereum_validators_monitoring_mentions{nos_global_index="1;1", mentions="op_1:ping_1,op_1:ping_2", alerts="UserSlashedValidators"}
      #   values: 1+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47
      - series: ethereum_validators_monitoring_user_operators_identifies{nos_name="Operator 1", nos_global_index="1;0"}
        values: 1+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47
      - series: ethereum_validators_monitoring_user_operators_identifies{nos_name="Operator 2", nos_global_index="1;1"}
        values: 1+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47
      - series: ethereum_validators_monitoring_user_operators_identifies{nos_name="Operator 3", nos_global_index="1;2"}
        values: 1+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47
      - series: ethereum_validators_monitoring_user_operators_identifies{nos_name="Operator 4", nos_global_index="1;3"}
        values: 1+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47
      - series: ethereum_validators_monitoring_user_operators_identifies{nos_name="Operator 5", nos_global_index="1;4"}
        values: 1+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47
      - series: ethereum_validators_monitoring_epoch_number
        values: 0+0x48 1+0x47 2+0x47 3+0x47 3+0x47 4+0x47 5+0x47 6+0x47
    alert_rule_test:
      - eval_time: 5m
        alertname: UserSlashedValidators
      - eval_time: 9m
        alertname: UserSlashedValidators
        exp_alerts:
          - exp_labels:
              severity: critical
              status: slashed
              nos_name: Operator 1
              mentions: op_0:ping_1,op_0:ping_2,op_0:ping_3
            exp_annotations:
              emoji: 🔪
              summary: "Operators have slashed validators"
              description: "Number of slashed validators per operator"
              field_name: "Operator 1"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"

          - exp_labels:
              severity: critical
              status: slashed
              nos_name: Operator 2
              mentions: op_1:ping_1,op_1:ping_2,op_1:ping_3,op_1:ping_4
            exp_annotations:
              emoji: 🔪
              summary: "Operators have slashed validators"
              description: "Number of slashed validators per operator"
              field_name: "Operator 2"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+2&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"

          - exp_labels:
              severity: critical
              status: slashed
              nos_name: Operator 3
              mentions: op_2:ping_3
            exp_annotations:
              emoji: 🔪
              summary: "Operators have slashed validators"
              description: "Number of slashed validators per operator"
              field_name: "Operator 3"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+3&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"

          - exp_labels:
              severity: critical
              status: slashed
              nos_name: Operator 4
            exp_annotations:
              emoji: 🔪
              summary: "Operators have slashed validators"
              description: "Number of slashed validators per operator"
              field_name: "Operator 4"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+4&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"

          - exp_labels:
              severity: critical
              status: slashed
              nos_name: Operator 5
            exp_annotations:
              emoji: 🔪
              summary: "Operators have slashed validators"
              description: "Number of slashed validators per operator"
              field_name: "Operator 5"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+5&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - eval_time: 14m
        alertname: UserSlashedValidators
      - eval_time: 17m
        alertname: UserSlashedValidators
        exp_alerts:
          - exp_labels:
              severity: critical
              status: slashed
              nos_name: Operator 1
              mentions: op_0:ping_1,op_0:ping_2,op_0:ping_3
            exp_annotations:
              emoji: 🔪
              summary: "Operators have slashed validators"
              description: "Number of slashed validators per operator"
              field_name: "Operator 1"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-180000.000000&to=1020000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 2'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
          - exp_labels:
              severity: critical
              status: slashed
              nos_name: Operator 2
              mentions: op_1:ping_1,op_1:ping_2,op_1:ping_3,op_1:ping_4
            exp_annotations:
              emoji: 🔪
              summary: "Operators have slashed validators"
              description: "Number of slashed validators per operator"
              field_name: "Operator 2"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+2&from=-180000.000000&to=1020000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 2'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - eval_time: 25m
        alertname: UserSlashedValidators
        exp_alerts:
          - exp_labels:
              severity: critical
              status: slashed
              nos_name: Operator 2
              mentions: op_1:ping_1,op_1:ping_2,op_1:ping_3,op_1:ping_4
            exp_annotations:
              emoji: 🔪
              summary: "Operators have slashed validators"
              description: "Number of slashed validators per operator"
              field_name: "Operator 2"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+2&from=300000.000000&to=1500000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 3'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - eval_time: 57m
        alertname: UserSlashedValidators
        exp_alerts:
          - exp_labels:
              severity: critical
              status: slashed
              nos_name: Operator 1
              mentions: op_0:ping_1,op_0:ping_2,op_0:ping_3
            exp_annotations:
              emoji: 🔪
              summary: "Operators have slashed validators"
              description: "Number of slashed validators per operator"
              field_name: "Operator 1"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=2220000.000000&to=3420000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 6'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"

  - interval: 10s
    input_series:
      - series: ethereum_validators_monitoring_data_actuality
        values: 3700000+0x180 3300000+0x180
    alert_rule_test:
      - alertname: DataActuality
        eval_time: 30m
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              emoji: ⏳
              summary: "Data actuality greater then 1 hour"
              description: "(1h 1m 40s) It's not OK. Please, check app health"
              resolved_summary: "Data actuality is back to normal and now less then 1 hour"
              resolved_description: "It's OK"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • '
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: DataActuality
        eval_time: 35m

  - interval: 10s
    input_series:
      - series: ethereum_validators_monitoring_validator_count_with_negative_balances_delta{nos_name="Operator 1"}
        values: 0+0x48 1+0x47 2+0x47 0+0x47
      - series: ethereum_validators_monitoring_user_operators_identifies{nos_name="Operator 1", nos_global_index="1;0"}
        values: 1+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47
      - series: ethereum_validators_monitoring_epoch_number
        values: 0+0x48 1+0x47 2+0x47 3+0x47 3+0x47 4+0x47
    alert_rule_test:
      - alertname: NumValidatorsWithNegativeDelta
        eval_time: 5m
      - alertname: NumValidatorsWithNegativeDelta
        eval_time: 9m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              mentions: op_0:ping_1,op_0:ping_2,
            exp_annotations:
              emoji: 💸
              summary: "Operators have a negative balance delta"
              description: "Number of validators per operator who have a negative balance delta."
              resolved_summary: "Operators have a positive balance delta"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: "Epoch • 1"
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithNegativeDelta
        eval_time: 17m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              mentions: op_0:ping_1,op_0:ping_2,
            exp_annotations:
              emoji: 💸
              summary: "Operators have a negative balance delta"
              description: "Number of validators per operator who have a negative balance delta."
              resolved_summary: "Operators have a positive balance delta"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-180000.000000&to=1020000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: "Epoch • 2"
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithNegativeDelta
        eval_time: 25m

  - interval: 10s
    input_series:
      - series: ethereum_validators_monitoring_validator_count_miss_attestation_last_n_epoch{nos_name="Operator 1", epoch_interval="3"}
        values: 0+0x48 1+0x47 2+0x47 2+0x47 2+0x47 0+0x47
      - series: ethereum_validators_monitoring_user_operators_identifies{nos_name="Operator 1", nos_global_index="1;0"}
        values: 1+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47
      - series: ethereum_validators_monitoring_epoch_number
        values: 0+0x48 1+0x47 2+0x47 3+0x47 3+0x47 4+0x47
    alert_rule_test:
      - alertname: NumValidatorsWithMissAttestationLastNEpoch
        eval_time: 5m
      - alertname: NumValidatorsWithMissAttestationLastNEpoch
        eval_time: 9m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
              mentions: op_0:ping_1,op_0:ping_2
            exp_annotations:
              emoji: 📝❌
              summary: "Operators have missed attestation in last 3 finalized epochs"
              description: "Number of validators per operator who have missed attestations."
              resolved_summary: "Operators not have missed attestation in last 3 finalized epochs"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithMissAttestationLastNEpoch
        eval_time: 14m
      - alertname: NumValidatorsWithMissAttestationLastNEpoch
        eval_time: 17m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
              mentions: op_0:ping_1,op_0:ping_2
            exp_annotations:
              emoji: 📝❌
              summary: "Operators have missed attestation in last 3 finalized epochs"
              description: "Number of validators per operator who have missed attestations."
              resolved_summary: "Operators not have missed attestation in last 3 finalized epochs"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-180000.000000&to=1020000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 2'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithMissAttestationLastNEpoch
        eval_time: 20m
      - alertname: NumValidatorsWithMissAttestationLastNEpoch
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
              mentions: op_0:ping_1,op_0:ping_2
            exp_annotations:
              emoji: 📝❌
              summary: "Operators have missed attestation in last 3 finalized epochs"
              description: "Number of validators per operator who have missed attestations."
              resolved_summary: "Operators not have missed attestation in last 3 finalized epochs"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=300000.000000&to=1500000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 3'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithMissAttestationLastNEpoch
        eval_time: 33m
      - alertname: NumValidatorsWithMissAttestationLastNEpoch
        eval_time: 41m

  - interval: 10s
    input_series:
      - series: ethereum_validators_monitoring_validator_count_high_inc_delay_last_n_epoch{nos_name="Operator 1", epoch_interval="3"}
        values: 0+0x48 1+0x47 2+0x47 2+0x47 2+0x47 0+0x47
      - series: ethereum_validators_monitoring_epoch_number
        values: 0+0x48 1+0x47 2+0x47 3+0x47 3+0x47 4+0x47
    alert_rule_test:
      - alertname: NumValidatorsWithHighIncDelayAttestationLastNEpoch
        eval_time: 5m
      - alertname: NumValidatorsWithHighIncDelayAttestationLastNEpoch
        eval_time: 9m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
            exp_annotations:
              emoji: 📝🐢
              summary: "Operators have attestation inc. delay greater than 2 in last 3 finalized epochs"
              description: "Number of validators per operator who have attestation with high inc. delay."
              field_name: "Operator 1"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithHighIncDelayAttestationLastNEpoch
        eval_time: 14m
      - alertname: NumValidatorsWithHighIncDelayAttestationLastNEpoch
        eval_time: 17m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
            exp_annotations:
              emoji: 📝🐢
              summary: "Operators have attestation inc. delay greater than 2 in last 3 finalized epochs"
              description: "Number of validators per operator who have attestation with high inc. delay."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-180000.000000&to=1020000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 2'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithHighIncDelayAttestationLastNEpoch
        eval_time: 20m
      - alertname: NumValidatorsWithHighIncDelayAttestationLastNEpoch
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
            exp_annotations:
              emoji: 📝🐢
              summary: "Operators have attestation inc. delay greater than 2 in last 3 finalized epochs"
              description: "Number of validators per operator who have attestation with high inc. delay."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=300000.000000&to=1500000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 3'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithHighIncDelayAttestationLastNEpoch
        eval_time: 33m
      - alertname: NumValidatorsWithHighIncDelayAttestationLastNEpoch
        eval_time: 41m

  - interval: 10s
    input_series:
      - series: ethereum_validators_monitoring_validator_count_invalid_attestation_property_last_n_epoch{nos_name="Operator 1", epoch_interval="3"}
        values: 0+0x48 1+0x47 2+0x47 2+0x47 2+0x47 0+0x47
      - series: ethereum_validators_monitoring_epoch_number
        values: 0+0x48 1+0x47 2+0x47 3+0x47 3+0x47 4+0x47
    alert_rule_test:
      - alertname: NumValidatorsWithInvalidPropertyAttestationLastNEpoch
        eval_time: 5m
      - alertname: NumValidatorsWithInvalidPropertyAttestationLastNEpoch
        eval_time: 9m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
            exp_annotations:
              emoji: '📝🏷️'
              summary: 'Operators have invalid attestation property in last 3 finalized epochs'
              description: 'Number of validators per operator who have invalid attestation property.'
              field_name: "Operator 1"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithInvalidPropertyAttestationLastNEpoch
        eval_time: 14m
      - alertname: NumValidatorsWithInvalidPropertyAttestationLastNEpoch
        eval_time: 17m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
            exp_annotations:
              emoji: '📝🏷️'
              summary: 'Operators have invalid attestation property in last 3 finalized epochs'
              description: 'Number of validators per operator who have invalid attestation property.'
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-180000.000000&to=1020000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 2'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithInvalidPropertyAttestationLastNEpoch
        eval_time: 20m
      - alertname: NumValidatorsWithInvalidPropertyAttestationLastNEpoch
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
            exp_annotations:
              emoji: '📝🏷️'
              summary: 'Operators have invalid attestation property in last 3 finalized epochs'
              description: 'Number of validators per operator who have invalid attestation property.'
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=300000.000000&to=1500000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 3'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithInvalidPropertyAttestationLastNEpoch
        eval_time: 33m
      - alertname: NumValidatorsWithInvalidPropertyAttestationLastNEpoch
        eval_time: 41m

  - interval: 10s
    input_series:
      - series: ethereum_validators_monitoring_validator_count_miss_propose{nos_name="Operator 1"}
        values: 0+0x48 1+0x47 2+0x47 2+0x47 2+0x47 0+0x47
      - series: ethereum_validators_monitoring_user_operators_identifies{nos_name="Operator 1", nos_global_index="1;0"}
        values: 1+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47
      - series: ethereum_validators_monitoring_epoch_number
        values: 0+0x48 1+0x47 2+0x47 3+0x47 3+0x47 4+0x47
    alert_rule_test:
      - alertname: NumValidatorsWithMissPropose
        eval_time: 5m
      - alertname: NumValidatorsWithMissPropose
        eval_time: 9m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              mentions: op_0:ping_1,op_0:ping_2
            exp_annotations:
              emoji: 📥
              summary: "Operators missed block propose in the last finalized epoch"
              description: "Number of validators per operator who missed block propose."
              resolved_summary: "Operators not missed block propose in the last finalized epoch"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithMissPropose
        eval_time: 14m
      - alertname: NumValidatorsWithMissPropose
        eval_time: 17m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              mentions: op_0:ping_1,op_0:ping_2
            exp_annotations:
              emoji: 📥
              summary: "Operators missed block propose in the last finalized epoch"
              description: "Number of validators per operator who missed block propose."
              resolved_summary: "Operators not missed block propose in the last finalized epoch"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-180000.000000&to=1020000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 2'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithMissPropose
        eval_time: 20m
      - alertname: NumValidatorsWithMissPropose
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              mentions: op_0:ping_1,op_0:ping_2
            exp_annotations:
              emoji: 📥
              summary: "Operators missed block propose in the last finalized epoch"
              description: "Number of validators per operator who missed block propose."
              resolved_summary: "Operators not missed block propose in the last finalized epoch"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=300000.000000&to=1500000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 3'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithMissPropose
        eval_time: 33m
      - alertname: NumValidatorsWithMissPropose
        eval_time: 41m

  - interval: 10s
    input_series:
      - series: ethereum_validators_monitoring_validator_count_with_sync_participation_less_avg_last_n_epoch{nos_name="Operator 1", epoch_interval="3"}
        values: 0+0x48 1+0x47 2+0x47 2+0x47 2+0x47 0+0x47
      - series: ethereum_validators_monitoring_user_operators_identifies{nos_name="Operator 1", nos_global_index="1;0"}
        values: 1+0x48 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47 1+0x47
      - series: ethereum_validators_monitoring_epoch_number
        values: 0+0x48 1+0x47 2+0x47 3+0x47 3+0x47 4+0x47
    alert_rule_test:
      - alertname: NumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 5m
      - alertname: NumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 9m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
              mentions: op_0:ping_1,op_0:ping_2
            exp_annotations:
              emoji: 🔄
              summary: "Operators sync participation less than average in last 3 finalized epochs"
              description: "Number of validators per operator whose sync participation less than average."
              resolved_summary: "Operators sync participation higher or equal than average in last 3 finalized epochs"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 14m
      - alertname: NumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 17m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
              mentions: op_0:ping_1,op_0:ping_2
            exp_annotations:
              emoji: 🔄
              summary: "Operators sync participation less than average in last 3 finalized epochs"
              description: "Number of validators per operator whose sync participation less than average."
              resolved_summary: "Operators sync participation higher or equal than average in last 3 finalized epochs"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-180000.000000&to=1020000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 2'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 20m
      - alertname: NumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
              mentions: op_0:ping_1,op_0:ping_2
            exp_annotations:
              emoji: 🔄
              summary: "Operators sync participation less than average in last 3 finalized epochs"
              description: "Number of validators per operator whose sync participation less than average."
              resolved_summary: "Operators sync participation higher or equal than average in last 3 finalized epochs"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=300000.000000&to=1500000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 3'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: NumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 33m
      - alertname: NumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 41m


  - interval: 10s
    input_series:
      - series: ethereum_validators_monitoring_high_reward_validator_count_with_sync_participation_less_avg_last_n_epoch{nos_name="Operator 1", epoch_interval="3"}
        values: 0+0x48 1+0x47 2+0x47 2+0x47 2+0x47 0+0x47
      - series: ethereum_validators_monitoring_epoch_number
        values: 0+0x48 1+0x47 2+0x47 3+0x47 3+0x47 4+0x47
    alert_rule_test:
      - alertname: HighRewardNumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 5m
      - alertname: HighRewardNumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 9m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
            exp_annotations:
              emoji: 📈🔄
              summary: "Operators may get high rewards in the future, but sync participation less than average in last 3 finalized epochs!"
              description: "Number of validators per operator whose sync participation less than average."
              resolved_summary: "Operators sync participation higher or equal than average in last 3 finalized epoch. Now may get high rewards in the future!"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: HighRewardNumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 14m
      - alertname: HighRewardNumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 17m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
            exp_annotations:
              emoji: 📈🔄
              summary: "Operators may get high rewards in the future, but sync participation less than average in last 3 finalized epochs!"
              description: "Number of validators per operator whose sync participation less than average."
              resolved_summary: "Operators sync participation higher or equal than average in last 3 finalized epoch. Now may get high rewards in the future!"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-180000.000000&to=1020000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 2'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: HighRewardNumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 20m
      - alertname: HighRewardNumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
            exp_annotations:
              emoji: 📈🔄
              summary: "Operators may get high rewards in the future, but sync participation less than average in last 3 finalized epochs!"
              description: "Number of validators per operator whose sync participation less than average."
              resolved_summary: "Operators sync participation higher or equal than average in last 3 finalized epoch. Now may get high rewards in the future!"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=300000.000000&to=1500000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 3'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: HighRewardNumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 33m
      - alertname: HighRewardNumValidatorsWithSyncParticipationLessAvgLastNEpoch
        eval_time: 41m

  - interval: 10s
    input_series:
      - series: ethereum_validators_monitoring_high_reward_validator_count_miss_attestation_last_n_epoch{nos_name="Operator 1", epoch_interval="3"}
        values: 0+0x48 1+0x47 2+0x47 2+0x47 2+0x47 0+0x47
      - series: ethereum_validators_monitoring_epoch_number
        values: 0+0x48 1+0x47 2+0x47 3+0x47 3+0x47 4+0x47
    alert_rule_test:
      - alertname: HighRewardNumValidatorsWithMissAttestationLastNEpoch
        eval_time: 5m
      - alertname: HighRewardNumValidatorsWithMissAttestationLastNEpoch
        eval_time: 9m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
            exp_annotations:
              emoji: 📈📝❌
              summary: "Operators may get high rewards in the future, but missed attestation in last 3 finalized epochs!"
              description: "Number of validators per operator who have missed attestations."
              resolved_summary: "Operators not have missed attestation in last 3 finalized epochs. Now may get high rewards in the future!"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: HighRewardNumValidatorsWithMissAttestationLastNEpoch
        eval_time: 14m
      - alertname: HighRewardNumValidatorsWithMissAttestationLastNEpoch
        eval_time: 17m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
            exp_annotations:
              emoji: 📈📝❌
              summary: "Operators may get high rewards in the future, but missed attestation in last 3 finalized epochs!"
              description: "Number of validators per operator who have missed attestations."
              resolved_summary: "Operators not have missed attestation in last 3 finalized epochs. Now may get high rewards in the future!"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-180000.000000&to=1020000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 2'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: HighRewardNumValidatorsWithMissAttestationLastNEpoch
        eval_time: 20m
      - alertname: HighRewardNumValidatorsWithMissAttestationLastNEpoch
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
              epoch_interval: 3
            exp_annotations:
              emoji: 📈📝❌
              summary: "Operators may get high rewards in the future, but missed attestation in last 3 finalized epochs!"
              description: "Number of validators per operator who have missed attestations."
              resolved_summary: "Operators not have missed attestation in last 3 finalized epochs. Now may get high rewards in the future!"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=300000.000000&to=1500000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 3'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: HighRewardNumValidatorsWithMissAttestationLastNEpoch
        eval_time: 33m
      - alertname: HighRewardNumValidatorsWithMissAttestationLastNEpoch
        eval_time: 41m

  - interval: 10s
    input_series:
      - series: ethereum_validators_monitoring_high_reward_validator_count_miss_propose{nos_name="Operator 1"}
        values: 0+0x48 1+0x47 2+0x47 2+0x47 2+0x47 0+0x47
      - series: ethereum_validators_monitoring_epoch_number
        values: 0+0x48 1+0x47 2+0x47 3+0x47 3+0x47 4+0x47
    alert_rule_test:
      - alertname: HighRewardNumValidatorsWithMissPropose
        eval_time: 5m
      - alertname: HighRewardNumValidatorsWithMissPropose
        eval_time: 9m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
            exp_annotations:
              emoji: 📈📥
              summary: "Operators may get high rewards in the future, but missed block propose in the last finalized epoch!"
              description: "Number of validators per operator who missed block propose."
              resolved_summary: "Operators not missed block propose in the last finalized epoch. Now may get high rewards in the future!"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[1](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-660000.000000&to=540000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 1'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: HighRewardNumValidatorsWithMissPropose
        eval_time: 14m
      - alertname: HighRewardNumValidatorsWithMissPropose
        eval_time: 17m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
            exp_annotations:
              emoji: 📈📥
              summary: "Operators may get high rewards in the future, but missed block propose in the last finalized epoch!"
              description: "Number of validators per operator who missed block propose."
              resolved_summary: "Operators not missed block propose in the last finalized epoch. Now may get high rewards in the future!"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=-180000.000000&to=1020000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 2'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: HighRewardNumValidatorsWithMissPropose
        eval_time: 20m
      - alertname: HighRewardNumValidatorsWithMissPropose
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              severity: critical
              nos_name: Operator 1
            exp_annotations:
              emoji: 📈📥
              summary: "Operators may get high rewards in the future, but missed block propose in the last finalized epoch!"
              description: "Number of validators per operator who missed block propose."
              resolved_summary: "Operators not missed block propose in the last finalized epoch. Now may get high rewards in the future!"
              resolved_description: "Number of validators per operator who recovered."
              field_name: "Operator 1"
              field_value: "[2](http://127.0.0.1:8082/d/3wimU2H7h/nodeoperators/?var-nos_name_var=Operator+1&from=300000.000000&to=1500000.000000)"
              url: "http://127.0.0.1:8082/d/HRgPmpNnz/validators"
              footer_text: 'Epoch • 3'
              footer_icon_url: "https://cryptologos.cc/logos/steth-steth-logo.png"
      - alertname: HighRewardNumValidatorsWithMissPropose
        eval_time: 33m
      - alertname: HighRewardNumValidatorsWithMissPropose
        eval_time: 41m
